```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Homework 2: Python Modules `re` and `numpy`

Welcome to DATA 271: Data Wrangling and Visualization! This is your second homework assignment of the semester. 

In this assignment, we will be practice using some common and powerful Python modules, the `re` module and the `numpy` module. These will help you prepare for the rest of the semester. As always, if you get stuck or have questions feel free to discuss problems with your peers and come to office hours. 

For this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `my_list` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you passed previously. **NOTE:** This homework assignment contains hidden tests. Passing all the auto-graded tests in the assignment does *NOT* mean you have answered everything correctly. So be careful and check your work!

```python
import otter
import re
import numpy as np
import matplotlib.pyplot as plt
grader = otter.Notebook()
```

## 1. Regex on Tweets

**Question 1.1.** The company that you are working for asked you to perform an analysis using a dataset with tweets. First of all, you need to do some cleaning and extract information.
While printing out some text, you realize that some tweets contain user mentions, and occasionally these mentions follow a very strange pattern. A few examples that you notice: `@robot3!`, `@robot5&` and `@robot7#`

To analyze if those users are bots, you will do a proof of concept with one tweet and extract them using the `.findall()` method.

The text of one tweet was saved in the variable `tweet`. Find all the matches in the `tweet` variable where the mentions start with `@`, followed by `robot` and have a number and a special character. e.g. `@robot3!`  **(5 Points)**

```python
tweet = "@robot9! @robot4& I have a good feeling that the show isgoing to be amazing! @robot9$ @robot7%"
bot_match = ...
bot_match
```

```python
grader.check("q1_1")
```

**Question 1.2.** While examining the tweet text in your dataset, you detect that some tweets carry more information. The text contains the number of retweets, user mentions, and likes. You decide to extract this important information that is given as in this example:

`this is what happens when you don't recycle your pizza boxes User_mentions:9, likes: 5, number of retweets: 4`

Write a regular expression that finds all user mentions info given as, for example, `User_mentions:9` in `tweet2` below.
**(5 Points)**

```python
tweet2 = "I pronounce tuple too-pull on Mon/Wed/Fri and tub-pull on Tue/Thu/Sat. On Sunday I don't talk about them. User_mentions:12, likes: 128, number of retweets: 75"
mentions = ...
mentions
```

```python
grader.check("q1_2")
```

**Question 1.3.** Write a regex that finds the number of likes in `tweet2` above. **(5 Points)**

```python
likes = ...
likes
```

```python
grader.check("q1_3")
```

**Question 1.4.** Write a regex that finds the number of retweets in `tweet2` above.
**(5 Points)**

```python
num_tweets = ...
num_tweets
```

```python
grader.check("q1_4")
```

**Question 1.5.** Some of the tweets in your dataset were downloaded incorrectly. Instead of having spaces to separate words, they have strange characters. You decide to use regular expressions to handle this situation. You print some of these tweets to understand which pattern you need to match.

You notice that the sentences are always separated by a special character, followed by a number, the word "break", and after that, another special character, e.g `&4break!`. The words are always separated by a special character, the word "new", and a normal random character, e.g `#newH`.

Write regular expressions to clean this data. First remove the strange sentence breaks in `clean_sentence_breaks`. Then replace the strange word breaks with a space in `clean_tweet`.
**(10 Points)**

```python
tweet3 = "He#newHis%newTdoing great in$newPscience. #8break%I am&newYvery impressed@newLalready."

clean_sentence_breaks = ...
print(clean_sentence_breaks)

clean_tweet = ...
print(clean_tweet)
```

```python
grader.check("q1_5")
```

**Question 6.** Sometimes there are sections of tweets that add complexity and messiness without adding useful information. For example, links and user mentions clutter up the tweets but aren't important for your company's analysis.

In order to clean the tweets, you want to extract some examples first. You know that most of the times links start with `http` and do not contain any whitespace, e.g. https://www.humboldt.edu. User mentions start with `@` and can have letters and numbers only, e.g. `@janedoe44`. Run the cell below to get import a list of tweets, which is a subset of the dataset you need to analyze. Then write regular expressions to create two lists of lists:
- `links` should be a list containing 10 lists. Each sublist cooresponds to a tweet in `list_of_tweets` and should contains the all links that appear in the tweet (if any).  
- `mentions_in_tweets` should be a list containing 10 lists. Each sublist cooresponds to a tweet in `list_of_tweets` and should contains the all mentions that appear in the tweet (if any).  

*NOTE:* Feel free to solve this with either a for loop or with list comprehension.
**(10 Points)**

```python
from csv import reader #Python's built-in csv module
# Open and read the datasett
tweets = open('tweets.csv')
tweets = reader(tweets) # type is csv.reader
list_of_tweets = [tweet[0] for tweet in list(tweets)]
list_of_tweets
```

```python
links = ...
print(links)

mentions_in_tweets = ...
print(mentions_in_tweets)
```

```python
grader.check("q1_6")
```

## 2. Regex on dates

Pretend you are working with an environmental organization monitoring water quality off the coast of California. Your ultimate goal is to observe how water quality changes through time. The data you've been given are collected by voluteer citizens who go out and take samples of the water occasionally, but unfortunately all of the volunteers record dates in different styles. This makes it difficult to analyze the data through time, so you must first clean the data. 

Run the following cell to import the data. (These data were fabricated for the purpose of this exercise). 

```python
# Open and read the datasett
water = open('water_data_dates.csv')
water = reader(water) # type is csv.reader
water = [i[0] for i in list(water)]
water[0]=re.sub("\ufeff","",water[0])
water
```

**Question 2.1.** Fill in the regular expression to find all the long-style dates in `water`. For example, find all the dates in the form `January 2 2023` or `January 02 2023` or `Jan 2 2023` or `Jan 02 2023.` or `Jan 2nd 2023.` (The specific month, day, and year do not have to match these examples, just the form). 
**(5 Points)**

```python
regex = ...
long_dates = [re.findall(regex,date)[0] for date in water if re.findall(regex,date) !=[]]
long_dates
```

```python
grader.check("q2_1")
```

**Question 2.2.** Fill in the regular expression to find all the short-style dates in `water`. For example, find all the dates in the form `1/2/23` or `01/02/2023` or `01/02/23` or `1/2/2023` or `1-1-2023` or `01-01-2023` or `01-01-23`  or `1-1-23`. (The specific month, day, and year do not have to match these examples, just the form). 
**(5 Points)**

```python
regex2 = ...
short_dates = [re.findall(regex2,date)[0] for date in water if re.findall(regex2,date) !=[]] 
short_dates
```

```python
grader.check("q2_2")
```

**Question 2.3.** Start cleaning the dates in `long_dates` to put them all in the form of `MON D YEAR`. For example `January 2 2023`, `January 2nd 2023` should become `Jan 2 2023`. (Again, the specific month, day, and year do not have to match this example, just the form). Your final answer should be a list which is the same length as `long_dates`.
**(5 Points)**

```python
clean_long_dates = ...
clean_long_dates
```

```python
grader.check("q2_3")
```

## 3. Numpy for Financial Analysis
Below is the daily stock prices of the S&P 500 for the 3 weeks.

```python
# Stock prices
stock_prices = np.array([4888.91,4886.66,4888.56,4856.8,4853.42,4796.28,4760.1,4739.13,4772.35,
4791.18,4792.13,4759.94,4741.93,4703.7,4690.57,4697.42,4725.07,4745.2,4782.88,4786.44,4773.45])
```

**Question 3.1.** Calculate the daily percent change in stock prices and find the standard deviation ofthe percent change. Recall that percent change is 

$$\frac{\text{new_value-old_value}}{\text{old_value}}\times100$$

*HINT:* You might find `np.diff` helpful. 
**(5 Points)**

```python
# Calculate daily percentage change
percent_change = ...

# Find standard deviation
sd_change = ...
sd_change
```

```python
grader.check("q3_1")
```

**Question 3.2:** *Trend analysis:* A moving average is a statistical technique used to smooth out fluctuations or noise in a dataset, making it easier to identify trends or patterns over time. It calculates the average of a subset of consecutive data points within a specified window or interval, and then moves the window forward one data point at a time, recalculating the average each time.

Compute the 5-day moving average of stock prices to identify long-term trends in stock prices. Put your results in a numpy array called `moving_average`.

*Hint:* To compute a moving average:

- Set the window size: 5 in this case
- Calculate the average: Take the average of the data points within the window at each step. This average becomes the value of the moving average at that position.
- Move the window: Slide the window forward by one data point and recalculate the average. Repeat this process until the window reaches the end of the dataset.
 
**(5 Points)**

```python
window_size = ...
moving_average = ...

...
    moving_average = ...

moving_average
```

```python
grader.check("q3_2")
```

**Question 3.3:** Now recompute the percent change from your moving average and find the standard deviation. What do you observe? **(5 Points)**

```python
# Calculate percentage change
percent_change_ma = ...

# Find standard deviation
sd_change_ma = ...
sd_change_ma
```

```python
grader.check("q3_3")
```

## 4. Numpy for Student Performance
Suppose you have scores for students across multiple exams in a class. Each row represents a student, and each column represents an exam. (These data are fabricated for the purpose of this exercise).

```python
num_students = 50
num_exams = 4
exam_scores = np.random.normal(82, 7, size=(num_students, num_exams))
```

**Question 4.1:** Compute the average score for each exams. Your answer should be an array with each entry being the average score of each exam. 
**(5 Points)**

```python
average_scores = ...
average_scores
```

```python
grader.check("q4_1")
```

**Question 4.2:** The first three columns are midterms worth 20% each, and the last column is the final worth 40% of the overall grade. Compute each student's overall grade in the class. 
**(5 Points)**

```python
student_grades = ...
student_grades
```

```python
grader.check("q4_2")
```

**Question 4.3:** How did the student who got the highest grade overall do on exam 2?
**(5 Points)**

```python
exam_2_grade = ...
exam_2_grade
```

```python
grader.check("q4_3")
```

## 5. Random Functions in Numpy for Simulation: Estimating $\pi$

In this problem we will use Numpy's random number generation to estimate $\pi$.

Consider the unit circle inscribed in a square. The circle, with radius $1$ has an area of $A = \pi$.  The square with sides of length $2$ has an area of $4$. The ratio of the area of the circle to the area of the square is therefore $\frac{\pi}{4}$.  If we throw darts randomly, the number of darts that land  in the circle divided by the total number of darts thrown will be described by the ratio of $\frac{\pi}{4}$.

It is enough to just consider the quarter circle and quarter square in the first quadrant, because the ratio in one quadrant should be the same as the ratio for the whole inscribed circle. Because the equation of the unit circle is $x^2+y^2 = 1$ we know a point is within the quarter circle if $\sqrt{x^2 + y^2} <1 $. Otherwise, it falls outside of the circle.

**Question 5.1:** Generate an array of size $1000 \times 2$ with random numbers between 0 and 1. The first column will represent the x-coordinates of 1000 randomly thrown darts, and the second column will be the y-coordinates. **(5 Points)**

```python
random_darts = ...
```

```python
grader.check("q5_1")
```

**Question 5.2:** Create an array with all the points from `random_darts` that fall inside the first quadrant of the unit circle. **(5 Points)**

```python
inside = ...
inside
```

```python
grader.check("q5_2")
```

**Question 5.3:** Use your previous two answers to get an approximation of $\pi$. **(5 Points)**

```python
pi_approx = ...
pi_approx
```

```python
grader.check("q5_3")
```

Run the following cell to visualize where your random darts landed. The blue ones are inside the unit circle, the red ones are outside. 

```python
# To visualize
plt.figure(figsize = (8, 8))
plt.scatter(random_darts[:, 0], random_darts[:, 1], s= 1, c = 'red') 
plt.scatter(inside[:, 0], inside[:, 1], s = 1, c = 'blue')
plt.show()
```

## You're done!

Congratulations! You finished Homework 2! Download the zip file and upload to Canvas. 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0])
```
>>> bot_match != ...
True
```


Tests for q1_2 (points: [0])
```
>>> mentions != ...
True
```


Tests for q1_3 (points: [0])
```
>>> likes != ...
True
```


Tests for q1_4 (points: [0])
```
>>> num_tweets != ...
True
```


Tests for q1_5 (points: [0])
```
>>> clean_sentence_breaks != ... and clean_tweet != ... 
True
```


Tests for q1_6 (points: [0])
```
>>> links != ... and mentions_in_tweets != ...
True
```


Tests for q2_1 (points: [0])
```
>>> regex != ...
True
```


Tests for q2_2 (points: [0])
```
>>> regex2 != ...
True
```


Tests for q2_3 (points: [0])
```
>>> clean_long_dates != ...
True
```


Tests for q3_1 (points: [0])
```
>>> len(percent_change) == 20
True
```


Tests for q3_2 (points: [0, 0])
```
>>> window_size != ...
True
```
```
>>> len(moving_average)
17
```


Tests for q3_3 (points: [0])
```
>>> len(percent_change_ma)
16
```


Tests for q4_1 (points: [0])
```
>>> len(average_scores)
4
```


Tests for q4_2 (points: [0])
```
>>> len(student_grades)
50
```


Tests for q4_3 (points: [0])
```
>>> exam_2_grade != ...
True
```


Tests for q5_1 (points: [0])
```
>>> random_darts.dtype == np.array([0.1]).dtype
True
```


Tests for q5_2 (points: [0])
```
>>> 750<= len(inside) <=790
True
```


Tests for q5_3 (points: [0])
```
>>> 3 < pi_approx < 3.3
True
```
